# Devoir maison 9

## Exercice 1

Le dénivelé cumulé positif d'une course de montagne est la somme totale des dénivelés de l'ensemble des ascensions durant la course. 

![exemple profil course](exemple.png)

Sur l'exemple ci-dessus :

* la course commence par une ascension de dénivelé positif $160$ ($490-330$) ;
* entre l'étape 2 et l'étape 3, le dénivelé positif est de $230$ ($610-380$) ;
* entre l'étape 3 et l'étape 4, le dénivelé positif est de $170$ ($780-610$) ;
* les autres parties de la course sont des descentes.

Le dénivelé cumulé positif total de cette course est donc $160+230+170=560$.

Écrire une fonction `denivele_positif` qui prend en argument la liste non vide des altitudes atteintes à la fin de chaque ascension et de chaque descente pendant la course et qui renvoie le dénivelé cumulé positif de cette course.

<!-- grade_id: cell-206fd07e779ab5a5, points: 0, grade: false, solution: true -->
```python
def denivele_positif(altitudes):
    ### BEGIN SOLUTION
    denivele = 0
    for i in range(len(altitudes) - 1):
        if altitudes[i] < altitudes[i+1]:
            denivele += altitudes[i+1] - altitudes[i]
    return denivele
    ### END SOLUTION
```

<!-- grade_id: cell-3261100124068e3d, points: 3, grade: true, solution: false -->
```python
assert denivele_positif([330, 490, 380, 610, 780, 550]) == 560
assert denivele_positif([200, 300, 100]) == 100
assert denivele_positif([150]) == 0
### BEGIN HIDDEN TESTS
assert denivele_positif([200]) == 0
assert denivele_positif([0, 1]) == 1
assert denivele_positif([1, 0]) == 0
assert denivele_positif([1, 0, 1]) == 1
assert denivele_positif([1, 2, 1]) == 1
assert denivele_positif([1, 1, 1]) == 0
assert denivele_positif([1, 2, 3]) == 2
assert denivele_positif([3, 2, 1]) == 0
### END HIDDEN TESTS
```

## Exercice 2

On se donne une chaîne de caractères appelée `motif` ainsi qu'un ensemble de règles du type "*Si le caractère lu est un* `a` *, remplace-le par* `ab`".

L'ensemble de ces règles est stocké dans un dictionnaire Python dans lequel :

* une clé est un caractère lu ;
* et la valeur associée est une chaine de caractères par laquelle le caractère doit être remplacé.

Si le caractère lu ne fait pas partie des clés du dictionnaire, il n'est associé à aucune règle et est donc recopié à l'identique.

Par exemple, si on prend pour règles `{'a': 'ab', 'b': 'ac', 'c': 'd'}` et pour motif de départ `'a'`, après **une** transformation on obtient `'ab'`.

Il est aussi possible d'effectuer cette transformation plusieurs fois de suite en l'appliquant à chaque étape au résultat de l'étape précédente. Par exemple :

- On garde pour règles `#!py {'a': 'ab', 'b': 'ac', 'c': 'd'}`.
- On prend pour encore motif de départ `'a'`.
- Après **une** transformation on obtient `'ab'`.
- Après **deux** transformations on obtient `'abac'`.  
- Après **trois** transformations on obtient `'abacabd'`.

Vous devez écrire deux fonctions Python :

* `transformation(motif, regles)` prend en argument un motif initial (une chaine de caractères) ainsi qu'un ensemble de règles (un dictionnaire) et renvoie la chaine obtenue après **une** transformation.

* `n_transformations(motif, regles, n)` prend en argument un motif initial (une chaine de caractères), un ensemble de règles (un dictionnaire) ainsi qu'un entier `n` et renvoie le résultat obtenu après `n` transformations.

<!-- grade_id: cell-65fff29dbc2e715c, points: 0, grade: false, solution: true -->
```python
def transformation(motif, regles):
    ### BEGIN SOLUTION
    resultat = ""
    for lettre in motif:
        if lettre in regles:
            resultat = resultat + regles[lettre]
        else:
            resultat = resultat + lettre
    return resultat
    ### END SOLUTION
```

<!-- grade_id: cell-8d30042acdbe1ee7, points: 2, grade: true, solution: false -->
```python
regles = {'a': 'ab', 'b': 'ac', 'c': 'd'}
motif = 'a'
assert transformation(motif, regles) == 'ab'

regles = {'a': 'aa'}
motif = 'a'
assert transformation(motif, regles) == 'aa'

### BEGIN HIDDEN TESTS
regles = {'a': 'b', 'b': 'a'}
motif = 'ac'
assert transformation(motif, regles) == 'bc'
### END HIDDEN TESTS
```

<!-- grade_id: cell-60a8514aaea5597f, points: 0, grade: false, solution: true -->
```python
def n_transformations(motif, regles, n):
    ### BEGIN SOLUTION
    if n == 0:
        resultat = motif
    else:
        resultat = transformation(motif, regles)
        for _ in range(n-1):
            resultat = transformation(resultat, regles)
    return resultat
    ### END SOLUTION
```

<!-- grade_id: cell-8d30042acdbe1ee8, points: 2, grade: true, solution: false -->
```python
regles = {'a': 'ab', 'b': 'ac', 'c': 'd'}
motif = 'a'
assert n_transformations(motif, regles, 2) == 'abac'
assert n_transformations(motif, regles, 3) == 'abacabd'
assert n_transformations("rien ne change !", {'z': 'y'}, 50) == 'rien ne change !'


### BEGIN HIDDEN TESTS
regles = {'a': 'a', 'b': 'b'}
motif = 'ab'
assert n_transformations(motif, regles, 20) == 'ab'
regles = {'a': 'aa'}
motif = 'a'
assert n_transformations(motif, regles, 4) == 'aaaaaaaaaaaaaaaa'
regles = {'a': 'b', 'b': 'a'}
motif = 'ac'
assert n_transformations(motif, regles, 4) == 'ac'
assert n_transformations(motif, regles, 5) == 'bc'
assert n_transformations(motif, regles, 0) == motif
### END HIDDEN TESTS
```

## Exercice 3

<h3 style="color:Tomato;"> Il est interdit d'utiliser la fonction <code>max</code> dans cet exercice.</h2>

Après un devoir, un professeur désire connaître les prénoms de l'élève ou des élèves ayant obtenu la meilleure note.

Les notes sont fournies dans un dictionnaire Python dans lequel les clés sont les prénoms des élèves et les valeurs leur note. Par exemple :

```python
notes = {'Alice': 18,
         'Bob': 10,
         'Charles': 8,
         'David': 12,
         'Erwann': 18,
         'Fanny': 14,
         'Guillaume': 18,
         'Hugo': 14}
```

On garantit que ce dictionnaire n'est pas vide et que toutes les notes sont positives ou nulles.

On désire créer une fonction `palmares` qui prend en paramètre un tel dictionnaire `notes` et renvoie un tuple de deux éléments :

* la note maximale qui a été attribuée ;

* et une liste contenant les noms du ou des élèves ayant obtenu cette note.


**Remarque : il est inutile de se soucier de l'ordre des prénoms dans la liste des résultats. Seule leur présence au sein de la liste importe.**

<!-- grade_id: cell-45e3f1184eb72f6f, points: 0, grade: false, solution: true -->
```python
def palmares(notes):
    ### BEGIN SOLUTION
    meilleure_note = 0
    eleves_meilleure_note = []
    for eleve in notes:
        if notes[eleve] > meilleure_note:
            meilleure_note = notes[eleve]
            eleves_meilleure_note = [eleve]
        elif notes[eleve] == meilleure_note:
            eleves_meilleure_note.append(eleve)

    return (meilleure_note, eleves_meilleure_note)
    ### END SOLUTION
```

<!-- grade_id: cell-b2bb5e71e244d930, points: 3, grade: true, solution: false -->
```python
notes = {'Alice': 18,
         'Bob': 10,
         'Charles': 8,
         'David': 12,
         'Erwann': 17,
         'Fanny': 14,
         'Guillaume': 16,
         'Hugo': 14}
assert palmares(notes) == (18, ['Alice'])

notes = {'Bob': 10,
         'Charles': 8,
         'David': 12,
         'Hugo': 14,
         'Fanny': 14}
assert palmares(notes)[0] == 14
assert sorted(palmares(notes)[1]) == ['Fanny', 'Hugo']

### BEGIN HIDDEN TESTS
notes = {'Romuald': 15,
         'Ronan': 14,
         'David': 15}
assert palmares(notes)[0] == 15
assert sorted(palmares(notes)[1]) == ['David', 'Romuald']
### END HIDDEN TESTS
```