# Devoir maison 14 &ndash; Liste des différences

Un laboratoire reçoit des résultats depuis deux sources différentes. Les résultats sont deux tableaux d'entiers `source_1` et `source_2` de même longueur.

Quand deux résultats concordent, on considère qu'il n'y aura pas à refaire l'expérience, on le codera par `False`, mais si les résultats sont différents, on considère qu'il faudra refaire cette expérience, on le codera par `True`.

Écrire une fonction telle que `differences(source_1, source_2)` renvoie un tableau de la longueur commune à `source_1` et `source_2` rempli de booléens.

<!-- grade_id: cell-3151ba1ebf2dbeba, points: 0, grade: false, solution: true -->
```python
def differences(source_1, source_2):
    ### BEGIN SOLUTION
    diff = []
    for i in range(len(source_1)):
        diff.append(source_1[i] != source_2[i])
    return diff
    ### END SOLUTION
```

<!-- grade_id: cell-9b93785a7be82ebe, points: 4, grade: true, solution: false -->
```python
assert differences([14, 87, 22, 5, 65], [14, 86, 27, 5, 65]) == [False, True, True, False, False]
assert differences([-54], [-54]) == [False]
assert differences([7, 8], [7, 11]) == [False, True]
assert differences([], []) == []
### BEGIN HIDDEN TESTS
source_1 = list(range(20))
source_2 = list(range(20))
assert differences(source_1, source_2) == [False] * 20

source_1 = list(range(20))
source_2 = list(range(1, 21))
assert differences(source_1, source_2) == [True] * 20

source_1 = list(range(21))
source_2 = list(range(20, -1, -1))
assert differences(source_1, source_2) == [True] * 10 + [False] + [True] * 10

for i in range(20):
    source_1 = list(range(20))
    source_2 = list(range(20))
    source_2[i] = -1
    attendu = [False] * 20
    attendu[i] = True
    assert differences(source_1, source_2) == attendu
### END HIDDEN TESTS
```