# Devoir maison 6

On dispose d'une liste, supposée non vide, de `0` et de `1` qui représente un nombre entier en base $2$, le bit de poids fort se trouvant à gauche.

On désire créer une fonction `ote_zeros(bins)` qui renvoie la liste obtenue en enlevant les éventuels zéros se trouvant à gauche de `bins`, c'est-à-dire les &laquo; zéros inutiles &raquo; de l'écriture en base $2$.

Si `bins` est une liste ne contenant que des `0`, c'est qu'elle représente l'entier zéro et la liste renvoyée sera `[0]`.

<!-- grade_id: cell-b8939708d3a3736b, points: 0, grade: false, solution: true -->
```python
def ote_zeros(bins):
    ### BEGIN SOLUTION
    resultat = []
    premier_1 = False
    for c in bins:
        if c == 1:
            premier_1 = True
        if premier_1:
            resultat.append(c)
    if not premier_1:
        resultat = [0]
    return resultat
    ### END SOLUTION
```

<!-- grade_id: cell-54e729700df23d1e, points: 2, grade: true, solution: false -->
```python
assert ote_zeros([0, 0, 0, 0,1, 0, 1]) == [1, 0, 1]
assert ote_zeros([0, 0, 0, 0]) == [0]
assert ote_zeros([1]) == [1]
### BEGIN HIDDEN TESTS
assert ote_zeros([0, 1, 0, 0, 0]) == [1, 0, 0, 0]
assert ote_zeros([1, 1, 1, 1, 0]) == [1, 1, 1, 1, 0]
assert ote_zeros([0, 0, 0, 1]) == [1]
### END HIDDEN TESTS
```