# Devoir maison 12 &ndash; Aplatir une liste

## Introduction

On considère un tableau à deux dimensions (on dit aussi une matrice), non vide, rempli de nombres et l'on souhaite l'*aplatir*, c'est-à-dire le transformer en un tableau à une seule ligne.

Par exemple, le tableau ci-dessous a pour dimensions $3 \times 4$ :

$$
\begin{array}{|c|c|c|c|}
\hline
1&2&3&4\\
\hline
5&6&7&8\\
\hline
9&10&11&12\\
\hline
\end{array}
$$

et pourra être transformé en un nouveau tableau de dimensions $1 \times 12$ ($1$ ligne et $12$ colonnes) :

$$
\begin{array}{|c|c|c|c|c|c|c|c|c|c|c|c|}
\hline
1&2&3&4&5&6&7&8&9&10&11&12\\
\hline
\end{array}
$$

Les tableaux *de départ* seront représentés par des listes de listes Python. Ainsi le premier tableau ci-dessus sera représenté par :

```python
matrice = [[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12]]
```

Les tableaux "aplatis", sur une seule ligne, seront représentés par une simple liste Python. Le second tableau ci-dessus sera représenté par :

```python
aplati = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12]
```

## Travail à faire
Écrire la fonction `aplatir` prenant en argument une liste de listes `tableau` et renvoyant une liste dans laquelle toutes les valeurs de `tableau` sont données à la suite les unes des autres.

<!-- grade_id: cell-201d3c2ca036973d, points: 0, grade: false, solution: true -->
```python
def aplatir(tableau):
    ### BEGIN SOLUTION
    tableau_aplati = []
    for sous_tableau in tableau:
        for element in sous_tableau:
            tableau_aplati.append(element)
    return tableau_aplati
    ### END SOLUTION
```

<!-- grade_id: cell-388a5f7bb7170f7c, points: 4, grade: true, solution: false -->
```python
assert aplatir([[1, 2, 3, 4], [5, 6, 7, 8]]) == [1, 2, 3, 4, 5, 6, 7, 8]
assert aplatir([[1], [2], [3], [4], [5], [6]]) == [1, 2, 3, 4, 5, 6]
### BEGIN HIDDEN TESTS
assert aplatir([[1, 2], [3, 4], [5, 6]]) == [1, 2, 3, 4, 5, 6]
assert aplatir([[1], [2], [3], [4], [5], [6]]) == [1, 2, 3, 4, 5, 6]
assert aplatir([[8, 7], [6, 5], [4, 3]]) == [8, 7, 6, 5, 4, 3]
assert aplatir([[1], [2], [3]]) == [1, 2, 3]
assert aplatir([[1, 1], [2, 2], [3, 3]]) == [1, 1, 2, 2, 3, 3]
assert aplatir([[1, 1]]) == [1, 1]
### END HIDDEN TESTS
```