# Devoir maison 8 &ndash; Adresse électronique

### Introduction

Une adresse électronique est composée de deux parties :

* la partie *locale*, avant le « @ »,
* la partie *serveur*, après le « @ ».

Certains serveurs de messagerie ignorent les caractères `.` dans la partie locale. Ainsi, sur ces serveurs, les messages envoyés aux adresses `johnconnor@hollywood.biz`, `john.connor@hollywood.biz`, `jo.hn.co.nn.or@hollywood.biz` et `john..connor@hollywood.biz` arrivent tous au même destinataire `johnconnor@hollywood.biz`. Par contre les adresses `john_connor@hollywood.biz` et `john_connore@hollywood.biz` ne correspondent pas à ce destinataire.

L'adresse de référence est l'adresse électronique écrite **sans aucun point dans la partie locale** : dans l'exemple précédent, l'adresse de référence est donc `johnconnor@hollywood.biz`.

### But de l'exercice

On cherche dans cet exercice à vérifier qu'une adresse correspond bien à une adresse de référence.

C'est le cas si, sans tenir compte des caractères `.` dans la partie locale, les deux adresses comportent les mêmes caractères présents dans le même ordre.

Dans cet exercice :

- On garantit que les parties « locales » des deux adresses sont non vides.

- On garantit que les parties « serveur » des deux adresses sont identiques.

### À vous de jouer !

1. Écrire la fonction `partie_locale` qui :
- prend en paramètre une adresse électronique `adresse` sous forme de chaîne de caractères ;
- et renvoie la partie locale sous forme de chaîne de caractères.

<!-- grade_id: cell-3324c718bdf84543, points: 0, grade: false, solution: true -->
```python
def partie_locale(adresse):
    ### BEGIN SOLUTION
    locale = ''
    i = 0
    while adresse[i] != '@':
        locale = locale + adresse[i]
        i = i + 1
    return locale
    ### END SOLUTION
```

<!-- grade_id: cell-e91aedd841f33835, points: 2, grade: true, solution: false -->
```python
assert partie_locale("johnconnor@hollywood.biz") == "johnconnor"
assert partie_locale("jo.hn.co.nn.or@hollywood.biz") == "jo.hn.co.nn.or"
### BEGIN HIDDEN TESTS
assert partie_locale("john..connor@hollywood.biz") == "john..connor"
assert partie_locale("john_connor@hollywood.biz") == "john_connor"
assert partie_locale("johnconnore@hollywood.biz") == "johnconnore"
### END HIDDEN TESTS
```

2. Écrire la fonction `partie_locale_sans_points` qui :
- prend en paramètre une adresse électronique `adresse` sous forme de chaîne de caractères ;
- et renvoie la partie locale sans points, sous forme de chaîne de caractères.

<!-- grade_id: cell-5dfbae8789231e3d, points: 0, grade: false, solution: true -->
```python
def partie_locale_sans_points(adresse):
    ### BEGIN SOLUTION
    locale_sans_points = ""
    for caractere in partie_locale(adresse):
        if caractere != ".":
            locale_sans_points += caractere
    return locale_sans_points
    ### END SOLUTION
```

<!-- grade_id: cell-e23d2840c13ce6dd, points: 2, grade: true, solution: false -->
```python
assert partie_locale_sans_points("johnconnor@hollywood.biz") == "johnconnor"
assert partie_locale_sans_points("jo.hn.co.nn.or@hollywood.biz") == "johnconnor"
assert partie_locale_sans_points("john..connor@hollywood.biz") == "johnconnor"
### BEGIN HIDDEN TESTS
assert partie_locale_sans_points("john_connor@hollywood.biz") == "john_connor"
assert partie_locale_sans_points("johnconnore@hollywood.biz") == "johnconnore"
### END HIDDEN TESTS
```

3. Écrire la fonction `correspond` qui prend en argument les deux adresses `adresse` et `reference` et renvoie `True` ou `False` selon que les deux adresses correspondent ou non.

<!-- grade_id: cell-2f3f2dd5573dcc6a, points: 0, grade: false, solution: true -->
```python
def correspond(adresse, reference):
    ### BEGIN SOLUTION
    return partie_locale_sans_points(adresse) == partie_locale(reference)
    ### END SOLUTION
```

<!-- grade_id: cell-4e5cad424ce7584b, points: 2, grade: true, solution: false -->
```python
assert correspond("johnconnor@hollywood.biz", "johnconnor@hollywood.biz") == True
assert correspond("john.connor@hollywood.biz", "johnconnor@hollywood.biz") == True
assert correspond("jo.hn.co.nn.or@hollywood.biz", "johnconnor@hollywood.biz") == True
### BEGIN HIDDEN TESTS
assert correspond("john..connor@hollywood.biz", "johnconnor@hollywood.biz") == True
assert correspond("john_connor@hollywood.biz", "johnconnor@hollywood.biz") == False
assert correspond("johnconnore@hollywood.biz", "johnconnor@hollywood.biz") == False
### END HIDDEN TESTS
```