# Devoir maison 11 &ndash; Soleil couchant sur des bâtiments

## Introduction
Lorsque des bâtiments sont alignés, ils se font de l'ombre les uns les autres. Dans cet exercice, nous sommes au soleil couchant, les rayons du soleil sont donc supposés horizontaux.

![](soleil-couchant.png)

Le schéma ci-dessus illustre un soleil couchant qui éclaire 9 bâtiments, les rayons du soleil sont représentés par des flèches horizontales.

- Les bâtiments aux indices `0` et `3` reçoivent des rayons de soleil alors que les bâtiments aux indices `1` et `2` sont masqués.
- Les **4** bâtiments aux indices `[0, 3, 6, 9]` reçoivent des rayons de soleil sur au moins un étage et sont donc éclairés, alors que les autres ne le sont pas.
- Il n'y a pas de bâtiment à l'indice `4`.

## Travail à faire
Écrire une fonction `nb_batiments_eclaires` qui prend en argument la liste `hauteurs` des bâtiments et qui renvoie le nombre de bâtiments éclairés.

- La hauteur des bâtiments (en nombre d'étages) est donnée par une liste d'entiers positifs. Une hauteur de zéro étage signifie l'absence de bâtiment.
- Pour l'exemple ci-dessus, cette liste est `[2, 1, 2, 4, 0, 4, 5, 3, 5, 6]`.

<!-- grade_id: cell-d9be3de7ff143d7b, points: 0, grade: false, solution: true -->
```python
def nb_batiments_eclaires(hauteurs):
    ### BEGIN SOLUTION
    nb = 0
    hauteur_max = 0
    for h in hauteurs:
        if h > hauteur_max:
            nb += 1
            hauteur_max = h
    return nb
    ### END SOLUTION
```

<!-- grade_id: cell-4be417c1d01d74cd, points: 4, grade: true, solution: false -->
```python
assert nb_batiments_eclaires([2, 1, 2, 4, 0, 4, 5, 3, 5, 6]) == 4
assert nb_batiments_eclaires([0, 3, 1, 2]) == 1
### BEGIN HIDDEN TESTS
assert nb_batiments_eclaires([1, 3, 2, 0, 2, 4, 5, 9, 7]) == 5
assert nb_batiments_eclaires([1, 0, 2, 1, 0]) == 2
assert nb_batiments_eclaires([3, 0, 2, 2, 5, 3, 6]) == 3
assert nb_batiments_eclaires([3, 4, 9, 2, 5, 3, 6]) == 3
### END HIDDEN TESTS
```