# Devoir maison 4 &ndash; Dernière occurrence

Programmer la fonction `derniere_occurrence`, prenant en paramètre un tableau non vide d'entiers et un entier `cible`, et qui renvoie l'indice de la **dernière** occurrence de `cible`.

Si l'élément n'est pas présent, la fonction renvoie la longueur du tableau.

> On n'utilisera pas la fonction `index`.

<!-- grade_id: cell-3335c4262b36fde6, points: 0, grade: false, solution: true -->
```python
def derniere_occurrence(tableau, cible):
    indice_dernier = len(tableau)
    ### BEGIN SOLUTION
    for i in range(len(tableau)):
        if tableau[i] == cible:
            indice_dernier = i
    return indice_dernier
    ### END SOLUTION
```

<!-- grade_id: cell-47817dc45303fc14, points: 3, grade: true, solution: false -->
```python
# 1 est absent, on renvoie la longueur du tableau : 2.
assert derniere_occurrence([5, 3], 1) == 2
# 2 est présent à l'indice 0, on renvoie 0.
assert derniere_occurrence([2, 4], 2) == 0
# 2 est présent aux indices 0 et 3, on renvoie le dernier : 3.
assert derniere_occurrence([2, 3, 5, 2, 4], 2) == 3
### BEGIN HIDDEN TESTS
assert derniere_occurrence([1, 2, 3, 1, 4, 5], 1) == 3
assert derniere_occurrence([1, 2, 3, 4, 5], 6) == 5
### END HIDDEN TESTS
```