# Devoir maison 5 &ndash; Nombre de répétitions

Écrire une fonction python appelée `nb_repetitions` qui prend en paramètres un élément `element` et un tableau de `valeurs` et renvoie le nombre de fois où l’élément `element` apparaît dans le tableau `valeurs`.

> Il est **interdit** d'utiliser la méthode `count`.

<!-- grade_id: cell-fc38ea85cd24c015, points: 0, grade: false, solution: true -->
```python
def nb_repetitions(element, valeurs):
    ### BEGIN SOLUTION
    nb = 0
    for elt in valeurs:
        if elt == element:
            nb = nb + 1
    return nb
    ### END SOLUTION
```

<!-- grade_id: cell-c958bc41bcbab94c, points: 2, grade: true, solution: false -->
```python
assert nb_repetitions(5, [2, 5, 3, 5, 6, 9, 5]) == 3
assert nb_repetitions('A', ['B', 'A', 'B', 'A', 'R']) == 2
assert nb_repetitions(12, [1, 7, 21, 36, 44]) == 0
### BEGIN HIDDEN TESTS
assert nb_repetitions('N', ['N', 'S', 'I']) == 1
assert nb_repetitions(-1, [0, 1, -1, -1, 6]) == 2
assert nb_repetitions('f', ['f', 'a', 'r', 'f', 'e', 'l', 'u']) == 2
### END HIDDEN TESTS
```