# Devoir maison 2 &ndash; Fonctions

## &laquo;Âge humain&raquo;

L'équivalence entre l'âge d'un chien et l'âge humain est donné en procédant ainsi :
- lorsque le chien a deux ans ou moins, chaque année compte pour douze années &laquo;humaines&raquo; ;
- lorsque le chien a plus de deux ans, chaque année au-delà des deux premières compte pour six années &laquo;humaines&raquo;.

1. Écrire un calcul permettant de déterminer l'&laquo; âge humain &raquo; d'un chien de 5 ans.

<!-- grade_id: cell-98143302141ee027, points: 1, grade: true, solution: true -->
```python
# À vous de jouer !
2 * 12 + 3 * 6
```

2. Écrire une fonction ``age_humain`` qui prend en paramètre l'âge $n$ ($n$ est un entier) d'un chien et qui renvoie l'&laquo;âge humain&raquo; de ce chien.

<!-- grade_id: cell-77bc15a76b52800f, points: 0, grade: false, solution: true -->
```python
def age_humain(n):
    ### BEGIN SOLUTION
    if n <= 2:
        return 12 * n
    else:
        return 24 + (n - 2) * 6
    ### END SOLUTION
```

<!-- grade_id: cell-dc6862996c9c4c1f, points: 3, grade: true, solution: false -->
```python
assert age_humain(1) == 12
assert age_humain(2) == 24
assert age_humain(4) == 36
### BEGIN HIDDEN TESTS
assert age_humain(10) == 72
assert age_humain(8) == 60
### END HIDDEN TESTS
```

3. En utilisant la fonction précédente, déterminer l'&laquo; âge humain &raquo; d'un chien de 7 ans (écrire dans la cellule ci-dessous l'appel permettant de déterminer cet &laquo; âge humain &raquo;).

<!-- grade_id: cell-16d215f0d636b556, points: 1, grade: true, solution: true -->
```python
# À vous de jouer !
age_humain(7)
```