# Devoir maison 7 &ndash; Les amis de mes amis

On considère le schéma suivant où une flèche allant d'une personne Α vers une
personne B indique que la personne A « suit » la personne B sur son compte
*Immediam*. On dit alors que B est un *ami* de A.

Voici un exemple de réseau *Immediam* :

![exemple réseau](exemple_reseau.png)

On représente ce réseau *Immediam* en machine par un dictionnaire dans lequel :

- les clés sont les chaînes de caractères correspondant aux noms des personnes inscrites ;
  
- les valeurs associées sont des listes de chaînes de caractères, représentant les personnes suivies.


Écrire une fonction  `amis_d_amis(reseau, membre)` qui :

- prend en argument un dictionnaire `reseau` représentant un tel réseau *Immediam* et une chaîne de caractères
`membre` qui représente un membre du réseau ;
- et qui renvoie la liste des amis des amis de `membre`, à l'exclusion de lui-même et sans doublon. 

**Indication : on peut vérifier facilement qu'un élément n'appartient pas à une liste.**

```py
>>> sports = ['badminton', 'tennis', 'tennis de table', 'squash']
>>> 'beach-volley' not in sports
True
```

Pour vous aider, voici le code incomplet de la fonction :
```py
def amis_d_amis(reseau, membre):
    resultat = []
    for ... in reseau[...]:
        for ami_de_ami in reseau[...]:
            if ... != membre and ... not in resultat:
                resultat.append(ami_de_ami)
    return resultat
```

<!-- grade_id: cell-ecb34983c4006dce, points: 0, grade: false, solution: true -->
```python
# À vous de jouer !
def amis_d_amis(reseau, membre):
    ### BEGIN SOLUTION
    resultat = []
    for ami in reseau[membre]:
        for ami_de_ami in reseau[ami]:
            if ami_de_ami != membre and ami_de_ami not in resultat:
                resultat.append(ami_de_ami)
    return resultat
    ### END SOLUTION
```

<!-- grade_id: cell-d6edcaa3f0178324, points: 4, grade: true, solution: false -->
```python
immediam = {
            "Anna": ["Billy"],
            "Billy": ["Anna", "Eroll"],
            "Carl": ["Billy"],
            "Dora": ["Gaby"],
            "Eroll": ["Billy", "Dora", "Flynn", "Gaby"],
            "Flynn": ["Gaby"],
            "Gaby": ["Eroll"],
    }
assert sorted(amis_d_amis(immediam, 'Billy')) == ['Dora', 'Flynn', 'Gaby']
assert sorted(amis_d_amis(immediam, 'Eroll')) == ['Anna', 'Gaby']
assert sorted(amis_d_amis(immediam, 'Dora')) == ['Eroll']
### BEGIN HIDDEN TESTS
assert sorted(amis_d_amis(immediam, 'Flynn')) == ['Eroll']
assert sorted(amis_d_amis(immediam, 'Carl')) == ['Anna', 'Eroll']
assert sorted(amis_d_amis(immediam, 'Gaby')) == ['Billy', 'Dora', 'Flynn']
assert sorted(amis_d_amis(immediam, 'Anna')) == ['Eroll']
### END HIDDEN TESTS
```