# Devoir maison 10 &ndash; Anniversaire de chat

## Introduction

Lorsqu'un chat souhaite offrir une carte d'anniversaire à un ami chat, il lui offre une carte avec un chat tout mignon avec un texte sur son âge dans sa vie actuelle. Oui, les chats ont plusieurs vies... Ce qui fait de très nombreuses cartes possibles à créer.

La société _CatCard_ propose de personnaliser la carte avec un texte en anglais (`"en"`) ou en français (`"fr"`) qui dépend de l'`age` et du `num_vie` (des entiers strictement positifs). Si langue demandée n'est ni `"fr"`, ni `"en"`, la société _CatCard_ propose une réponse `'Je donne ma langue au chat'`.

### Exemples de cartes

![cartes en français et en anglais](channiv_fr_en.jpeg)

### Travail à faire

Écrire une fonction `channiv` qui renvoie le texte demandé suivant les exemples suivants :

```python
>>> channiv(3, 4)
'Joyeux 3e channiversaire de ta 4e vie'
>>> channiv(3, 4, "en")
'Happy purrthday, for 3rd year of 4th life'
>>> channiv(1, 1, "fr")
'Joyeux 1er channiversaire de ta 1re vie'
>>> channiv(1, 1, "en")
'Happy purrthday, for 1st year of 1st life'
>>> channiv(3, 4, "miaou")
'Je donne ma langue au chat'
```

#### Indication

On pourra constater que `channiv(3, 4)` est équivalent à `channiv(3, 4, "fr")`, en effet la majorité des clients de CatCard sont français et veulent que la fonction soit configurée pour renvoyer par défaut un texte en français.

**Le mécanisme pour ce faire est le paramètre par défaut `lang` dans la définition de `channiv` :**

    ```python
    def channiv(age, num_vie, lang="fr"):
    ```

#### Remarque

- En français, on place **seulement** un `e` après le nombre, sauf pour `1er` (premier) et `1re` (première).

    > On trouve souvent des abréviations en `ème`, elles ne sont pas recommandées, sauf pour le cas d'une lettre comme nième, ou pième.

- En anglais, on place `th` après le nombre, sauf pour `1st` (_first_), `2nd` (_second_) et `3rd` (_third_).

<!-- grade_id: cell-89b5d7080d117834, points: 0, grade: false, solution: true -->
```python
# À vous de jouer !
def channiv(age, num_vie, lang="fr"):
    ### BEGIN SOLUTION
    if lang == "fr":
        carte = "Joyeux "
        carte += str(age)
        if age == 1:
            carte += "er"
        else:
            carte += "e"
        carte += " channiversaire de ta "
        carte += str(num_vie)
        if num_vie == 1:
            carte += "re"
        else:
            carte += "e"
        carte += " vie"
        return carte
    elif lang == "en":
        carte = "Happy purrthday, for "
        carte += str(age)
        if age == 1:
            carte += "st"
        elif age == 2:
            carte += "nd"
        elif age == 3:
            carte += "rd"
        else:
            carte += "th"
        carte += " year of "
        carte += str(num_vie)
        if num_vie == 1:
            carte += "st"
        elif num_vie == 2:
            carte += "nd"
        elif num_vie == 3:
            carte += "rd"
        else:
            carte += "th"
        carte += " life"
        return carte
    else:
        return "Je donne ma langue au chat"
    ### END SOLUTION
```

<!-- grade_id: cell-b2d8fc6e330fac18, points: 5, grade: true, solution: false -->
```python
assert channiv(3, 4) == 'Joyeux 3e channiversaire de ta 4e vie'
assert channiv(3, 4, "en") == 'Happy purrthday, for 3rd year of 4th life'
assert channiv(1, 1, "fr") == 'Joyeux 1er channiversaire de ta 1re vie'
assert channiv(1, 1, "en") == 'Happy purrthday, for 1st year of 1st life'
assert channiv(3, 4, "miaou") == 'Je donne ma langue au chat'
### BEGIN HIDDEN TESTS
assert channiv(4, 1) == 'Joyeux 4e channiversaire de ta 1re vie'
assert channiv(4, 2, "en") == 'Happy purrthday, for 4th year of 2nd life'
assert channiv(2, 4, "fr") == 'Joyeux 2e channiversaire de ta 4e vie'
assert channiv(3, 2, "en") == 'Happy purrthday, for 3rd year of 2nd life'
assert channiv(2, 2, "miaaaouuu") == 'Je donne ma langue au chat'
### END HIDDEN TESTS
```