# Devoir maison 13 &ndash; Nombre de zéros à la fin d'un entier

On souhaite avoir une fonction `nb_zeros` qui détermine le nombre de zéros à la fin de l'écriture décimale d'un entier $n>0$, très grand.

🙅‍♂️ **On s'interdit**, ici, d'utiliser la fonction de conversion `str`. Cette méthode est totalement inefficace avec des nombres très grands.
    
On demande plutôt de compter combien de fois on peut diviser un nombre par $10$ avec un reste égal à zéro. Par exemple, 

    1. 42000 = 4200×10 + 0,
    2. 4200 = 420×10 + 0,
    3. 420 = 42×10 + 0,
    4. 42 n'est pas divisible par 10.

On a pu diviser 42000 **trois fois** par 10 avec un reste égal à 0. Ce nombre se finit donc par **3** zéros.

<!-- grade_id: cell-ea693e4cf8daa7ce, points: 0, grade: false, solution: true -->
```python
# À vous de jouer !
def nb_zeros(n):
    ### BEGIN SOLUTION
    # non valable pour n = 0
    resultat = 0
    while n % 10 == 0:
        n = n // 10
        resultat += 1
    return resultat
    ### END SOLUTION
```

<!-- grade_id: cell-0b253ce290242c94, points: 4, grade: true, solution: false -->
```python
assert nb_zeros(42000) == 3
assert nb_zeros(3210) == 1
assert nb_zeros(282475249) == 0
assert nb_zeros(7**10000) == 0
assert nb_zeros(7**10000 * 1000) == 3
### BEGIN HIDDEN TESTS
def NB_ZEROS(n):
    resultat = 0
    while n % 10 == 0:
        n = n // 10
        resultat += 1
    return resultat

for n in range(1, 123):
    attendu = NB_ZEROS(n)
    assert nb_zeros(n) == attendu, f"Erreur avec n = {n}"
### END HIDDEN TESTS
```