# Devoir maison 22 &ndash; La perfection d'un mot

## 1 Mot parfait

### 1.1 Définition

On affecte à chaque lettre de l'alphabet un code selon le tableau ci-dessous :

|`'A'`|`'B'`|`'C'`|`'D'`|`'E'`|`'F'`|`'G'`|`'H'`|`'I'`|
|:---:|:---:|:---:|:---:|:---:|:---:|:---:|:---:|:---:|
|$1$|$2$|$3$|$4$|$5$|$6$|$7$|$8$|$9$|

|`'J'`|`'K'`|`'L'`|`'M'`|`'N'`|`'O'`|`'P'`|`'Q'`|`'R'`|
|:---:|:---:|:---:|:---:|:---:|:---:|:---:|:---:|:---:|
|$10$|$11$|$12$|$13$|$14$|$15$|$16$|$17$|$18$|

|`'S'`|`'T'`|`'U'`|`'V'`|`'W'`|`'X'`|`'Y'`|`'Z'`|
|:---:|:---:|:---:|:---:|:---:|:---:|:---:|:---:|
|$19$|$20$|$21$|$22$|$23$|$24$|$25$|$26$|


Pour un mot donné (une chaine de caractères non vide uniquement composée de lettres majuscules), on détermine :

- d'une part <span style="color:blue">son code alphabétique concaténé</span>, obtenu par la juxtaposition du texte des codes de chacun de ses caractères, et lu comme un unique entier ;
- et d'autre part, <span style="color:blue">son code additionné</span>, qui est la somme des codes de chacun de ses caractères.

<p style="color:red"><b>Définition : on dit que ce mot est « parfait » si le code additionné divise le code concaténé.</b></p>

### 1.2 Exemples

1. Pour le mot `"PAUL"`, les codes sont $16, 1, 21, 12$.
    - Le code concaténé est la chaine `1612112`, soit l'entier $1\,612\,112$.
    - Son code additionné est l'entier $16 + 1 + 21 + 12$ qui donne $50$.
    - $50$ ne divise pas l'entier $1\,612\,112$ ; par conséquent, le mot `"PAUL"` n'est pas parfait.

2. Pour le mot `"ALAIN"`, les codes sont $1, 12, 1, 9, 14$.
    - Le code concaténé est la chaine `1121914`, soit l'entier $1\,121\,914$.
    - Le code additionné est l'entier $37$, car $1 + 12 + 1 + 9 + 14 = 37$.
    - $37$ divise l'entier $1\,121\,914$ ; par conséquent, le mot `"ALAIN"` est parfait.

## 2 Travail à faire

### 2.1 Rappels
Pour cet exercice on pourra utiliser `str` et `int` comme fonctions de conversion.
    
D'autre part, on pourra utiliser la fonction `ord` qui renvoie le code ASCII d'un caractère ASCII passé en paramètre.

```py
>>> ord('A')
65
>>> ord('B')
66
```

```py
>>> str(12)
'12'
>>> str(12) + str(14)
'1214'
>>> int('1214')
1214
>>> int('1214') % 10
4
```

### 2.2 À vous de jouer !
1. Compléter la fonction `code` qui prend comme argument une lettre (majuscule) et qui renvoie son code tel que défini par l'énoncé.
2. Compléter ensuite la fonction `est_parfait`qui prend comme argument une chaine de caractères `mot` (en lettres majuscules) et qui renvoie le code alphabétique concaténé, le code additionné de `mot`, ainsi qu'un booléen qui indique si `mot` est parfait ou pas.

Par exemple :

```py
>>> est_parfait("PAUL")
(50, 1612112, False)
>>> est_parfait("ALAIN")
(37, 1121914, True)
```

<!-- grade_id: cell-8ee5dee0dab86928, points: 0, grade: false, solution: true -->
```python
def code(lettre):
    ### BEGIN SOLUTION
    return ord(lettre) - ord('A') + 1
    ### END SOLUTION
```

<!-- grade_id: cell-d087bb4fac4f73f0, points: 2, grade: true, solution: false -->
```python
assert code('A') == 1
assert code('T') == 20
### BEGIN HIDDEN TESTS
assert code('U') == 21
assert code('L') ==  12
### END HIDDEN TESTS
```

<!-- grade_id: cell-5093bf13c555ad7f, points: 0, grade: false, solution: true -->
```python
def est_parfait(mot):
    code_int_cumul = 0
    code_str_concatene = ""
    for lettre in mot:
    ### BEGIN SOLUTION
        valeur = code(lettre)
        code_int_cumul += valeur
        code_str_concatene += str(valeur)
    code_concatene = int(code_str_concatene)
    mot_est_parfait = (code_concatene % code_int_cumul == 0)  # un booléen
    return (code_int_cumul, code_concatene, mot_est_parfait)
    ### END SOLUTION
```

<!-- grade_id: cell-b238f41da8c9f2c9, points: 3, grade: true, solution: false -->
```python
assert est_parfait("PAUL") == (50, 1612112, False)
assert est_parfait("ALAIN") == (37, 1121914, True)
### BEGIN HIDDEN TESTS
def EST_PARFAIT(mot) :
    code_int_cumul = 0
    code_str_concatene = ""
    for lettre in mot:
        valeur = ord(lettre) - ord('A') + 1
        code_int_cumul += valeur
        code_str_concatene += str(valeur)
    code_concatene = int(code_str_concatene)
    mot_est_parfait = (code_concatene % code_int_cumul == 0)  # un booléen
    return (code_int_cumul, code_concatene, mot_est_parfait)

for i in range(26):
    mot = chr(ord('A') + i)
    resultat = est_parfait(mot)
    attendu = EST_PARFAIT(mot)
    assert attendu == resultat, f"Erreur avec le mot {mot}"

for j in range(26):
    for j in range(26):
        mot = chr(ord('A') + i) + chr(ord('A') + j)
        resultat = est_parfait(mot)
        attendu = EST_PARFAIT(mot)
        assert attendu == resultat, f"Erreur avec le mot {mot}"
### END HIDDEN TESTS
```