# Devoir maison 18 &ndash; Adresse IP, masque, adresse de sous-réseau

## 1 Fonction dec_to_bin

<div style="background-color:Lavender;">
    <h2>Exercice 1</h2>
    <h4>Compléter la fonction <tt>dec_to_bin</tt> afin qu'elle respecte la spécification donnée.</h4>
</div>

```python
def dec_to_bin(n):
    """ Entrée : un entier n compris entre 0 et 255
        Sortie : une liste de longueur 8 contenant uniquement des 0 et des 1
        et donnant l'écriture en base 2 de l'entier n
        Exemple : dec_to_bin(125) renvoie la liste [0, 1, 1, 1, 1, 1, 0, 1].
    """
    bin_tab = []
    nb_bits = 0
    while nb_bits != ...:
        bin_tab = [...] + bin_tab
        n = ...
        nb_bits = nb_bits + 1
    return ...
```

<!-- grade_id: cell-ced69f87ed50401e, points: 0, grade: false, solution: true -->
```python
def dec_to_bin(n):
    """ Entrée : un entier n compris entre 0 et 255
        Sortie : une liste de longueur 8 contenant uniquement des 0 et des 1
        et donnant l'écriture en base 2 de l'entier n
        Exemple : dec_to_bin(125) renvoie la liste [0, 1, 1, 1, 1, 1, 0, 1].
    """
    ### BEGIN CORRECTION
    bin_tab = []
    nb_bits = 0
    while nb_bits != 8:
        bin_tab = [n % 2] + bin_tab
        n = n // 2
        nb_bits = nb_bits + 1
    return bin_tab
    ### END CORRECTION
```

<!-- grade_id: cell-6e4b08a46e82bd5e, points: 2, grade: true, solution: false -->
```python
assert dec_to_bin(125) == [0, 1, 1, 1, 1, 1, 0, 1]
assert dec_to_bin(92) == [0, 1, 0, 1, 1, 1, 0, 0]
assert dec_to_bin(0) == [0, 0, 0, 0, 0, 0, 0, 0]
assert dec_to_bin(255) == [1, 1, 1, 1, 1, 1, 1, 1]
### BEGIN HIDDEN TESTS
assert dec_to_bin(42) == [0, 0, 1, 0, 1, 0, 1, 0]
assert dec_to_bin(189) == [1, 0, 1, 1, 1, 1, 0, 1]
assert dec_to_bin(18) == [0, 0, 0, 1, 0, 0, 1, 0]
### END HIDDEN TESTS
```

<div style="background-color:Lavender;">
    <h2>Exercice 2</h2>
    <h4>Compléter la fonction <tt>et</tt> afin qu'elle respecte la spécification donnée.</h4>
</div>

<!-- grade_id: cell-0df67af1aaed316e, points: 0, grade: false, solution: true -->
```python
def et(bin_tab_1, bin_tab_2):
    """ Entrées : deux listes de longueur 8 contenant uniquement des 0 et des 1
        Sortie : une liste de longueur 8 obtenue en faisant l'opération 'ET' logique
        bit à bit entre bin_tab_1 et bin_tab_2
        Exemple : avec bin_tab_1 = [0, 0, 1, 1, 0, 0, 1, 1]
                  et   bin_tab_2 = [1, 1, 1, 0, 1, 1, 0, 1] et(bin_tab_1, bin_tab_2) 
                  renvoie          [0, 0, 1, 0, 0, 0, 0, 1]
    """
    et_tab = []
    ### BEGIN CORRECTION
    for i in range(8):
        if bin_tab_1[i] * bin_tab_2[i] == 1:
            et_tab.append(1)
        else:
            et_tab.append(0)
    return et_tab
    ### END CORRECTION
```

<!-- grade_id: cell-bf5cd21493a26b44, points: 2, grade: true, solution: false -->
```python
bin_tab_1 = [0, 0, 1, 1, 0, 0, 1, 1]
bin_tab_2 = [1, 1, 1, 0, 1, 1, 0, 1] 
assert et(bin_tab_1, bin_tab_2) == [0, 0, 1, 0, 0, 0, 0, 1]

bin_tab_1 = [0, 0, 0, 0, 0, 0, 0, 1]
bin_tab_2 = [1, 1, 1, 1, 0, 0, 0, 1]
assert et(bin_tab_1, bin_tab_2) == [0, 0, 0, 0, 0, 0, 0, 1]
### BEGIN HIDDEN TESTS
bin_tab_1 = [0, 1, 1, 1, 0, 0, 1, 1]
bin_tab_2 = [1, 1, 0, 0, 1, 1, 0, 1] 
assert et(bin_tab_1, bin_tab_2) == [0, 1, 0, 0, 0, 0, 0, 1]

bin_tab_1 = [0, 0, 1, 1, 1, 0, 0, 1]
bin_tab_2 = [0, 1, 1, 1, 0, 0, 0, 1]
assert et(bin_tab_1, bin_tab_2) == [0, 0, 1, 1, 0, 0, 0, 1]
### END HIDDEN TESTS
```

<div style="background-color:Lavender;">
    <h2>Exercice 3</h2>
    <h4>Compléter la fonction <tt>bin_to_dec</tt> afin qu'elle respecte la spécification donnée.</h4>
</div>

<!-- grade_id: cell-f6243bc5adf484c8, points: 0, grade: false, solution: true -->
```python
def bin_to_dec(bin_tab):
    """ Entrée : une liste de longueur 8 contenant uniquement des 0 et des 1
        Sortie : un entier n dont la représentation en base 2 est donnée par la liste bin_tab
        Exemple : bin_to_dec([0, 1, 1, 1, 1, 1, 0, 1]) renvoie 125.
    """
    ### BEGIN CORRECTION
    n = 0
    for i in range(8):
        n = n + bin_tab[i] * 2**(7-i)
    return n
    ### END CORRECTION
```

<!-- grade_id: cell-deb967b5e2d5caed, points: 2, grade: true, solution: false -->
```python
assert bin_to_dec([0, 1, 1, 1, 1, 1, 0, 1]) == 125
assert bin_to_dec([0, 1, 0, 1, 1, 1, 0, 0]) == 92
assert bin_to_dec([0, 0, 0, 0, 0, 0, 0, 0]) == 0
assert bin_to_dec([1, 1, 1, 1, 1, 1, 1, 1]) == 255
### BEGIN HIDDEN TESTS
assert bin_to_dec([0, 0, 1, 0, 1, 0, 1, 0]) == 42
assert bin_to_dec([1, 0, 1, 1, 1, 1, 0, 1]) == 189
assert bin_to_dec([0, 0, 0, 1, 0, 0, 1, 0]) == 18
### END HIDDEN TESTS
```

<div style="background-color:Lavender;">
    <h2>Exercice 4</h2>
    <h4>Compléter la fonction <tt>adresse_sous_reseau</tt> afin qu'elle respecte la spécification donnée.</h4>
</div>

<!-- grade_id: cell-1ab7282263d4e1e1, points: 0, grade: false, solution: true -->
```python
def adresse_sous_reseau(ip, masque):
    """ Entrées : deux listes ip, masque de longueur 4 contenant uniquement des nombres compris entre 0 et 255
        donnant l'adresse IP et le masque d'une machine M
        Sortie : une liste de longueur 4 contenant uniquement des nombres compris entre 0 et 255 donnant l'adresse
        du sous-réseau auquel appartient la machine M
        Exemple : avec   ip = [129, 175, 127, 1]
                  et masque = [255, 255, 0, 0] adresse_sous_reseau(ip, masque)
                  renvoie     [129, 175, 0, 0]
    """
    sous_reseau = []
    ### BEGIN CORRECTION
    for i in range(4):
        sous_reseau.append(bin_to_dec(et(dec_to_bin(ip[i]), dec_to_bin(masque[i]))))
    return sous_reseau
    ### END CORRECTION
```

<!-- grade_id: cell-f726871397934100, points: 2, grade: true, solution: false -->
```python
assert adresse_sous_reseau([129, 175, 127, 1], [255, 255, 0, 0]) == [129, 175, 0, 0]
assert adresse_sous_reseau([129, 175, 127, 1], [255, 255, 255, 0]) == [129, 175, 127, 0]
### BEGIN HIDDEN TESTS
assert adresse_sous_reseau([192, 168, 0, 1], [255, 255, 255, 0]) == [192, 168, 0, 0]
assert adresse_sous_reseau([192, 168, 0, 3], [255, 255, 255, 0]) == [192, 168, 0, 0]
assert adresse_sous_reseau([129, 175, 127, 1], [255, 255, 252, 0]) == [129, 175, 124, 0]
assert adresse_sous_reseau([129, 175, 130, 10], [255, 255, 252, 0]) == [129, 175, 128, 0]
### END HIDDEN TESTS
```

<div style="background-color:Lavender;">
    <h2>Exercice 5</h2>
    <h4>Compléter la fonction <tt>meme_sous_reseau</tt> afin qu'elle respecte la spécification donnée.</h4>
</div>

<!-- grade_id: cell-88710a547f0d5677, points: 0, grade: false, solution: true -->
```python
def meme_sous_reseau(ip_1, masque_1, ip_2, masque_2):
    """ Entrées : quatre listes ip_1, masque_1, ip_2, masque_2 de longueur 4 contenant des nombres de 0 à 255
        donnant respectivement l'adresse IP et le masque d'une machine M1, l'adresse IP et le masque d'une machine M2
        Sortie : un booléen précisant si M1 et M2 appartiennent à un même sous-réseau (True si c'est le cas, False sinon)

    """
    appartient = True
    ### BEGIN CORRECTION
    sous_reseau_1 = adresse_sous_reseau(ip_1, masque_1)
    sous_reseau_2 = adresse_sous_reseau(ip_2, masque_2)
    for i in range(4):
        if sous_reseau_1[i] != sous_reseau_2[i]:
            appartient = False
    return appartient   
    ### END CORRECTION
```

<!-- grade_id: cell-5e0d659658fdfc74, points: 2, grade: true, solution: false -->
```python
ip_1 = [192, 168, 0, 1]
masque_1 = [255, 255, 255, 0]
ip_2 = [192, 168, 0, 3]
masque_2 = [255, 255, 255, 0]
assert meme_sous_reseau(ip_1, masque_1, ip_2, masque_2) == True

ip_1 = [129, 175, 127, 1]
masque_1 = [255, 255, 252, 0]
ip_2 = [129, 175, 130, 10]
masque_2 = [255, 255, 252, 0]
assert meme_sous_reseau(ip_1, masque_1, ip_2, masque_2) == False

### BEGIN HIDDEN TESTS
ip_1 = [129, 175, 130, 10]
masque_1 = [255, 255, 252, 0]
ip_2 = [129, 175, 128, 17]
masque_2 = [255, 255, 252, 0]
assert meme_sous_reseau(ip_1, masque_1, ip_2, masque_2) == True

ip_1 = [129, 175, 128, 17]
masque_1 = [255, 255, 252, 0]
ip_2 = [129, 175, 131, 110]
masque_2 = [255, 255, 252, 0]
assert meme_sous_reseau(ip_1, masque_1, ip_2, masque_2) == True

ip_1 = [129, 175, 131, 110]
masque_1 = [255, 255, 252, 0]
ip_2 = [129, 175, 132, 8]
masque_2 = [255, 255, 252, 0]
assert meme_sous_reseau(ip_1, masque_1, ip_2, masque_2) == False
### END HIDDEN TESTS
```