# Devoir maison 15 &ndash; Vacances de Noël

## Exercice 1 &ndash; Tous différents

### Introduction
Un tableau peut contenir plusieurs fois le même élément. C'est le cas du tableau `tableau_1` ci-dessous :
```python
tableau_1 = [1, 2, 3, 6, 2, 4, 5]
```
La valeur `2` est deux fois dans ce tableau.

Au contraire, dans le tableau `tableau_2`, toutes les valeurs sont uniques :
```python
tableau_2 = ['chien', 'chat', 'lion', 'poisson']
```

### Travail à faire
Écrire une fonction `tous_differents` qui prend un tableau (qu'on supposera non vide) `tableau` et renvoie un booléen indiquant si toutes les valeurs de `tableau` sont différentes ou non.

**Indication.&ndash; On pourra utiliser `range(a, b)` qui itère sur toutes les valeurs entières allant de `a` inclus à `b` exclu.**

<!-- grade_id: cell-bff765ec4c5c63a3, points: 0, grade: false, solution: true -->
```python
# À vous de jouer !
def tous_differents(tableau):
### BEGIN SOLUTION
    resultat = True
    n = len(tableau)
    for i in range(n - 1):
        for j in range(i + 1, n):
            if tableau[i] == tableau[j]:
                resultat = False
    return resultat
### END SOLUTION
```

<!-- grade_id: cell-38cb49fcd54d8d69, points: 4, grade: true, solution: false -->
```python
tableau_1 = [1, 2, 3, 6, 2, 4, 5]
assert tous_differents(tableau_1) == False
tableau_2 = ['chien', 'chat', 'lion', 'poisson']
assert tous_differents(tableau_2) == True
### BEGIN HIDDEN TESTS
tableau_3 = [2, 4, 1, 0]
assert tous_differents(tableau_3) == True
tableau_4 = ['bleu', 'rouge', 'bleu', 'vert']
assert tous_differents(tableau_4) == False
### END HIDDEN TESTS
```

## Exercice 2 &ndash; Le bon enclos

### Introduction

On considère une liste de dictionnaires Python qui contient des enregistrements relatifs à des animaux hébergés dans un refuge.

Les attributs des enregistrements sont :

* `nom` ;
* `espece` ;
* `age` ;
* `enclos`.
 
Les valeurs associées à `nom` et `espece` sont des chaines de caractères, celles associées à `age` et `enclos` des entiers.

Voici un exemple d'une telle table :

```python
animaux = [ {'nom': 'Medor', 'espece': 'chien', 'age': 5, 'enclos': 2},
            {'nom': 'Titine', 'espece': 'chat', 'age': 2, 'enclos': 5},
            {'nom': 'Tom', 'espece': 'chat', 'age': 7, 'enclos': 4},
            {'nom': 'Belle', 'espece': 'chien', 'age': 6, 'enclos': 3},
            {'nom': 'Mirza', 'espece': 'chat', 'age': 6, 'enclos': 5}]
```

On garantit que chaque enregistrement contient l'ensemble des informations (aucune clé ne manque dans un dictionnaire).

### Travail à faire

Programmer une fonction `selection_enclos` qui :

* prend en paramètres :
    * une liste `table_animaux` contenant des enregistrements relatifs à des animaux (comme dans l'exemple ci-dessus) ;
    * un numéro d'enclos `num_enclos`,
* renvoie une table contenant les enregistrements de `table_animaux` dont l'attribut `enclos` est égal à `num_enclos`. Dans cette table, les enregistrements seront donnés dans le même ordre que dans la table d'entrée.


Exemples avec la table animaux ci-dessus :


```python
>>> selection_enclos(animaux, 5)
[{'nom': 'Titine', 'espece': 'chat', 'age': 2, 'enclos': 5},
{'nom':'Mirza', 'espece':'chat', 'age':6, 'enclos':5}]
>>> selection_enclos(animaux, 2)
[{'nom':'Medor', 'espece':'chien', 'age':5, 'enclos':2}]
>>> selection_enclos(animaux, 7)
[]
```

<!-- grade_id: cell-215b46efeea99479, points: 0, grade: false, solution: true -->
```python
# À vous de jouer !
def selection_enclos(table_animaux, num_enclos):
    ### BEGIN SOLUTION
    resultat = []
    for animal in table_animaux:
        if animal['enclos'] == num_enclos:
            resultat.append(animal)
    return resultat
    ### END SOLUTION
```

<!-- grade_id: cell-b34350e76f869fe6, points: 3, grade: true, solution: false -->
```python
animaux = [{'nom': 'Medor', 'espece': 'chien', 'age': 5, 'enclos': 2},
           {'nom': 'Titine', 'espece': 'chat', 'age': 2, 'enclos': 5},
           {'nom': 'Tom', 'espece': 'chat', 'age': 7, 'enclos': 4},
           {'nom': 'Belle', 'espece': 'chien', 'age': 6, 'enclos': 3},
           {'nom': 'Mirza', 'espece': 'chat', 'age': 6, 'enclos': 5}]
assert selection_enclos(animaux, 5) == [{'nom': 'Titine', 'espece': 'chat', 'age': 2, 'enclos': 5}, {
    'nom': 'Mirza', 'espece': 'chat', 'age': 6, 'enclos': 5}]
assert selection_enclos(animaux, 2) == [
    {'nom': 'Medor', 'espece': 'chien', 'age': 5, 'enclos': 2}]
assert selection_enclos(animaux, 7) == []
### BEGIN HIDDEN TESTS
enclos = {animal['enclos'] for animal in animaux}
for num_enclos in enclos:
    reponse = [ani for ani in animaux if ani['enclos'] == num_enclos]
    assert selection_enclos(animaux, num_enclos) == reponse

assert selection_enclos(animaux, -1) == []

animaux = [{'nom': 'Moder', 'espece': 'lapin', 'age': 5, 'enclos': 0},
           {'nom': 'Ollie', 'espece': 'furet', 'age': 6, 'enclos': 0}]
assert selection_enclos(animaux, 0) == animaux
assert selection_enclos(animaux, 1) == []
### END HIDDEN TESTS
```

## Exercice 3 &ndash; Mot qui correspond à un mot à trous

### Introduction

**Mot à trous**  
- Un mot est ici une chaine de caractères composée uniquement de lettres de l'alphabet.
- Un mot à trous comporte également zéro, une ou plusieurs fois le caractère `"."`.

On dira que `mot_complet` correspond à `mot_a_trous`, si on peut remplacer chaque `"."` de `mot_a_trous` par une lettre de façon à obtenir `mot_complet`. 

Par exemple :  
- `"INFO.MA.IQUE"` est un mot à trous ;
- `"INFORMATIQUE"` est un mot qui lui correspond ;
- `"AUTOMATIQUE"` est un mot qui ne lui correspond pas.

### Travail à faire

Écrire une fonction telle que `correspond(mot_complet, mot_a_trous)` renvoie un booléen qui détermine si `mot_complet` correspond à `mot_a_trous`.

<!-- grade_id: cell-646e281576034a8d, points: 0, grade: false, solution: true -->
```python
# À vous de jouer
def  correspond(mot_complet, mot_a_trous):
    ### BEGIN SOLUTION
    resultat = True
    if len(mot_complet) != len(mot_a_trous):
        resultat = False
    else:
        for i in range(len(mot_complet)):
            if mot_a_trous[i] != "." and mot_a_trous[i] != mot_complet[i]:
                resultat = False
    return resultat
    ### END SOLUTION
```

<!-- grade_id: cell-db8d92ca2623a5fe, points: 3, grade: true, solution: false -->
```python
assert correspond("MAL", "MAL..") == False
assert correspond("INFORMATIQUE", "INFO.MA.IQUE") == True
assert correspond("AUTOMATIQUE", "INFO.MA.IQUE") == False
assert correspond("INFO", "INFO.MA.IQUE") == False
assert correspond("INFORMATIQUES", "INFO.MA.IQUE") == False
### BEGIN HIDDEN TESTS
assert correspond("MOTO", "M.T.") == True
assert correspond("ACCIDENT", "..CIDENT") == True
assert correspond("ACCIDENT", "I.CIDENT") == False
assert correspond("CONJECTURE", "CONJO.CTURE") == False
### END HIDDEN TESTS
```