# Devoir maison 17 &ndash; Refuge d'animaux

Le refuge Wouaf_Miaou s'occupe de chats et de chiens pouvant être adoptés.
Compléter la fonction `selection_enclos` afin de renvoyer la liste de noms d'animaux d'un même enclos. L'ordre de rangement des noms dans la liste n'a pas d'importance.

Par exemple :

```py
>>> Wouaf_Miaou = [
        {'nom':'Medor', 'espece':'chien', 'age':5, 'enclos':2},
        {'nom':'Titine', 'espece':'chat', 'age':2, 'enclos':5},
        {'nom':'Tom', 'espece':'chat', 'age':7, 'enclos':4},
        {'nom':'Belle', 'espece':'chien', 'age':6, 'enclos':3},
        {'nom':'Mirza', 'espece':'chat', 'age':6, 'enclos':5},
]
>>> sorted(selection_enclos(Wouaf_Miaou, 5))
['Mirza', 'Titine']
>>> selection_enclos(Wouaf_Miaou, 2)
['Medor']
>>> selection_enclos(Wouaf_Miaou, 7)
[]
```

**Remarque : l'ordre de parcours d'un dictionnaire dans Python se fait dans un ordre arbitraire. On utilise `sorted` dans les tests pour que les éléments de la réponse soient forcément dans l'ordre alphabétique et ainsi pouvoir les comparer avec les résultats attendus.**

<!-- grade_id: cell-4432621567eaacf4, points: 0, grade: false, solution: true -->
```python
# À vous de jouer !
def selection_enclos(refuge, numero_enclos):
    ### BEGIN SOLUTION
    noms = []
    for animal in refuge:
        if animal['enclos'] == numero_enclos:
            noms.append(animal['nom'])
    return noms
    ### END SOLUTION
```

<!-- grade_id: cell-f4d1866a18139ad2, points: 3, grade: true, solution: false -->
```python
Wouaf_Miaou = [
    {'nom':'Medor', 'espece':'chien', 'age':5, 'enclos':2},
    {'nom':'Titine', 'espece':'chat', 'age':2, 'enclos':5},
    {'nom':'Tom', 'espece':'chat', 'age':7, 'enclos':4},
    {'nom':'Belle', 'espece':'chien', 'age':6, 'enclos':3},
    {'nom':'Mirza', 'espece':'chat', 'age':6, 'enclos':5},
]

assert sorted(selection_enclos(Wouaf_Miaou, 5)) == ['Mirza', 'Titine']
assert selection_enclos(Wouaf_Miaou, 2) == ['Medor']
assert selection_enclos(Wouaf_Miaou, 7) == []

# BEGIN HIDDEN TESTS
Wouaf_Miaou = [
    {'nom':'Medor', 'espece':'chien', 'age':5, 'enclos':2},
    {'nom':'Titine', 'espece':'chat', 'age':2, 'enclos':4},
    {'nom':'Tom', 'espece':'chat', 'age':7, 'enclos':4},
    {'nom':'Belle', 'espece':'chien', 'age':6, 'enclos':3},
    {'nom':'Mirza', 'espece':'chat', 'age':6, 'enclos':5},
    {'nom':'Capucine', 'espece':'chat', 'age':3, 'enclos':3},
    {'nom':'Oreo', 'espece':'chat', 'age':4, 'enclos':4}
]

assert sorted(selection_enclos(Wouaf_Miaou, 3)) == ['Belle', 'Capucine']
assert sorted(selection_enclos(Wouaf_Miaou, 4)) == ['Oreo', 'Titine', 'Tom']
# END HIDDEN TESTS
```