# Mot qui correspond à un mot à trous

## Introduction

**Mot à trous**  
- Un mot est ici une chaine de caractères composée uniquement de lettres de l'alphabet.
- Un mot à trous comporte également zéro, une ou plusieurs fois le caractère `"."`.

On dira que `mot_complet` correspond à `mot_a_trous`, si on peut remplacer chaque `"."` de `mot_a_trous` par une lettre de façon à obtenir `mot_complet`. 

Par exemple :  
- `"INFO.MA.IQUE"` est un mot à trous ;
- `"INFORMATIQUE"` est un mot qui lui correspond ;
- `"AUTOMATIQUE"` est un mot qui ne lui correspond pas.

## Travail à faire

Écrire une fonction telle que `correspond(mot_complet, mot_a_trous)` renvoie un booléen qui détermine si `mot_complet` correspond à `mot_a_trous`.

<!-- grade_id: cell-7862e2447a80203f, points: 0, grade: false, solution: true -->
```python
# À vous de jouer
def  correspond(mot_complet, mot_a_trous):
    ### BEGIN SOLUTION
    resultat = True
    if len(mot_complet) != len(mot_a_trous):
        resultat = False
    else:
        for i in range(len(mot_complet)):
            if mot_a_trous[i] != "." and mot_a_trous[i] != mot_complet[i]:
                resultat = False
    return resultat
    ### END SOLUTION
```

<!-- grade_id: cell-e71e6d23d5cfbdaa, points: 4, grade: true, solution: false -->
```python
assert correspond("MAL", "MAL..") == False
assert correspond("INFORMATIQUE", "INFO.MA.IQUE") == True
assert correspond("AUTOMATIQUE", "INFO.MA.IQUE") == False
assert correspond("INFO", "INFO.MA.IQUE") == False
assert correspond("INFORMATIQUES", "INFO.MA.IQUE") == False
### BEGIN HIDDEN TESTS
assert correspond("MOTO", "M.T.") == True
assert correspond("ACCIDENT", "..CIDENT") == True
assert correspond("ACCIDENT", "I.CIDENT") == False
assert correspond("CONJECTURE", "CONJO.CTURE") == False
### END HIDDEN TESTS
```