# Devoir maison 19 &ndash; Chiffres romains

## 1 Introduction

Les chiffres romains sont un système ancien d'écriture des nombres.

Les chiffres (symboles) romains sont: I, V, X, L, C, D et M. Ces symboles représentent respectivement 1, 5, 10, 50, 100, 500 et 1000 en base dix.

Cette association pourra être modélisée par un dictionnaire défini une fois pour toute (une constante) donné ci-dessous :

```py
VALEUR = {"I": 1, "V": 5, "X": 10, "L": 50, "C": 100, "D": 500, "M": 1000}
```

## 2 Principe de l'algorithme

Le principe est de commencer à zéro, puis d'additionner ou retrancher les valeurs de tous les symboles représentant un nombre écrit en chiffres romains. 

* Si un symbole a une valeur **supérieure ou égale** à celui qui se trouve à sa droite, il est **ajouté**. Par exemple :

  "XVI" est le nombre 16  car "X" vaut 10 qui est supérieur à 5 (valeur de "V"), qui est lui même supérieur à 1 (valeur de "I").  
    
  Ainsi "XVI" vaut 0 + 10 + 5 + 1 = 16. 
    

* Si un symbole a une valeur **strictement inférieure** à celui qui se trouve à sa droite, il est **retranché**. Par exemple :

  "XIV" est le nombre 14  car "X" vaut 10 qui est supérieur ou égal à 1 (valeur de "I"), qui est lui même inférieur à 5 (valeur de "V"). 
  
  Ainsi "XIV" vaut 0 + 10 - 1 + 5 = 14.  


* Si le symbole est en dernière position, il est **ajouté**.

## 3 À vous de jouer !

On souhaite créer une fonction `valeur` qui prend en paramètre une chaîne de caractères (non vide) représentant un nombre écrit en chiffres romains et qui renvoie sa valeur en écriture décimale.

```py
>>> valeur("XVI")
16
>>> valeur("MMXXII")
2022
>>> valeur("CDII")
402
>>> valeur("XLII")
42
```  

<!-- grade_id: cell-73e00af4bcfccdf0, points: 0, grade: false, solution: true -->
```python
# À vous de jouer !
def valeur(romain):
    VALEUR = {"I": 1, "V": 5, "X": 10, "L": 50, "C": 100, "D": 500, "M": 1000}
    ### BEGIN SOLUTION
    n = len(romain)
    cumul = 0
    for i in range(n - 1):
        if VALEUR[romain[i]] < VALEUR[romain[i + 1]]:
            cumul = cumul - VALEUR[romain[i]]
        else:
            cumul = cumul + VALEUR[romain[i]]
    cumul = cumul + VALEUR[romain[n - 1]]
    return cumul    
    ### END SOLUTION
```

<!-- grade_id: cell-5c637c9ea08faf6f, points: 4, grade: true, solution: false -->
```python
assert valeur("XVI") == 16
assert valeur("MMXXII") == 2022
assert valeur("CDII") == 402
assert valeur("XLII") == 42
### BEGIN HIDDEN TESTS
assert valeur("MCMXCI") == 1991
assert valeur("MCMXCIX") == 1999
assert valeur("MMMCMXCIX") == 3999
assert valeur("XCIV") == 94
assert valeur("MDCLXVI") == 1666
### END HIDDEN TESTS
```