# Devoir maison 21 &ndash; Suite de Syracuse d'un entier

## 1 Introduction

Pour construire la suite de Syracuse d'un entier $n$ :

1. On utilise une variable éponyme initialisée à $n$.
2. On crée une liste avec $n$ comme premier élément,
3. On répète en boucle tant que $n \neq 1$ :
    - Si $n$ est pair, on le divise par $2$,
    - Sinon $n$ devient $3n+1$.
    - On ajoute $n$ à la liste.

On admet que, quel que soit l'entier $n$ choisi au départ, la suite finit toujours par atteindre la valeur $1$.

Écrire une fonction `syracuse` prenant en paramètres un entier `n` strictement positif et qui renvoie la suite de Syracuse de l'entier `n`.

> On rappelle que `n % 2` renvoie le reste dans la division de `n` par `2`. Si ce reste est nul, c'est que `n` est un nombre pair.

Exemples :

```py
>>> syracuse(7)
[7, 22, 11, 34, 17, 52, 26, 13, 40, 20, 10, 5, 16, 8, 4, 2, 1]
>>> syracuse(3)
[3, 10, 5, 16, 8, 4, 2, 1]
>>> syracuse(1)
[1]
```

## 2 À vous de jouer !

<!-- grade_id: cell-3585a864004808b8, points: 0, grade: false, solution: true -->
```python
def syracuse(n):
    ### BEGIN SOLUTION
    suite = [n]
    while n != 1:
        if n % 2 == 0:
            n = n // 2
        else:
            n = 3 * n + 1
        suite.append(n)
    return suite
    ### END SOLUTION
```

<!-- grade_id: cell-9c74e4cb1677d878, points: 4, grade: true, solution: false -->
```python
assert syracuse(7) == [7, 22, 11, 34, 17, 52, 26, 13, 40, 20, 10, 5, 16, 8, 4, 2, 1]
assert syracuse(3) == [3, 10, 5, 16, 8, 4, 2, 1]
assert syracuse(1) == [1]
### BEGIN HIDDEN TESTS
def solution(n):
    suite = [n]
    while n != 1:
        if n % 2 == 0:
            n = n // 2
        else:
            n = 3 * n + 1
        suite.append(n)
    return suite

for n in range(1, 200):
    resultat = solution(n)
    assert syracuse(n) == resultat, f"Erreur avec {n}"
### END HIDDEN TESTS
```