# Devoir maison 20 &ndash; Gelées

## 1 Introduction

Un météorologue a relevé les températures au lever du jour dans sa rue.
Il souhaite déterminer la durée de la plus longue période de gelées consécutives durant ces relevés.

Vous devez écrire une fonction `gelees(temperatures)` qui renvoie la longueur de la plus longue séquence de nombres négatifs ou nuls consécutifs dans la liste.

Les températures sont données sous forme d'une liste de nombres :

```py
temperatures = [2, -3, -2, 0, 1, -1]
```

Si la liste est vide, la fonction renverra la valeur `0`.

On rappelle que l'eau gèle à partir de 0°C inclus.

```py
>>> gelees([2, -3, -2, 0, 1, -1])
3
>>> gelees([3, 2, 2])
0
>>> gelees([])
0
```

## 2 À vous de jouer !

<!-- grade_id: cell-87e67b57d8ac3eaa, points: 0, grade: false, solution: true -->
```python
def gelees(temperatures):
    ### BEGIN SOLUTION
    plus_longue_periode = 0
    periode = 0
    for temp in temperatures:
        if temp <= 0:
            periode += 1
            if periode > plus_longue_periode:
                plus_longue_periode = periode
        else:
            periode = 0
    return plus_longue_periode
    ### END SOLUTION
```

<!-- grade_id: cell-513252cd80e30414, points: 4, grade: true, solution: false -->
```python
assert gelees([2, -3, -2, 0, 1, -1]) == 3
assert gelees([3, 2, 2]) == 0
assert gelees([]) == 0
### BEGIN HIDDEN TESTS
assert gelees([-2, -3, -2, 0, -1, -1]) == 6
assert gelees([0, 0, 0]) == 3
assert gelees([1]*1000+[-1]+[1]*1000) == 1
### END HIDDEN TESTS
```