# Devoir maison 23 &ndash; Vacances de Pâques

## Exercice 1 &ndash; Pangramme

#### Définition et exemples

Un *pangramme* est une phrase contenant toutes les lettres de l'alphabet.

Il existe dans l'alphabet latin utilisé en français 26 lettres minuscules « usuelles ». Si l'on tient compte des lettres accentuées et des ligatures (*à*, *â*, ..., *ç*, *æ* et *œ*), le décompte passe à 42 caractères minuscules distincts.

**On se limite dans cet exercice aux lettres minuscules usuelles de l'alphabet : *a*, *b*, *c*, ..., *z*.**

- Un exemple classique de pangramme est « *portez ce vieux whisky au juge blond qui fume* ». Chaque lettre y apparaît **au moins une fois**.
- Par contre « *portez __un__ vieux whisky au juge blond qui fume* » n'est pas un pangramme ; il manque la lettre *c*.

#### Travail à faire

Vous devez écrire une fonction `est_pangramme` qui :

* prend en paramètre une chaîne de caractères `phrase`,
* renvoie le booléen indiquant si cette chaîne est un pangramme ou non.

Pour vous aider, on fournit la fonction indice qui prend en argument une lettre entre `"a"` et `"z"` et renvoie son indice dans l'alphabet (en débutant à `0` pour `"a"`). **Il n'est pas indispensable de l'utiliser.**

```python
def indice(minuscule):
    return ord(minuscule) - ord("a")
```

<!-- grade_id: cell-d68e0612bd86cd1d, points: 0, grade: false, solution: true -->
```python
def  est_pangramme(phrase):
    ### BEGIN SOLUTION
    presences = [False] * 26
    
    for caractere in phrase:
        if "a" <= caractere <= "z":
            presences[indice(caractere)] = True

    for i in range(26):
        if not presences[i]:
            return False

    return True
    ### END SOLUTION
```

<!-- grade_id: cell-f2313430b317d837, points: 4, grade: true, solution: false -->
```python
assert est_pangramme("portez ce vieux whisky au juge blond qui fume !") == True
assert est_pangramme("portez un vieux whisky au juge blond qui fume !") == False
assert est_pangramme("jugez que ce texte renferme l'alphabet, dix voyelles, k et w") == True
assert est_pangramme("jugez que ce texte renferme l'alphabet, dix voyelles et w") == False
### BEGIN HIDDEN TESTS
from string import ascii_lowercase
assert est_pangramme("the quick brown fox jumps over The lazy dog.") == True
assert est_pangramme("hier, au zoo, j'ai vu dix guépards, cinq zébus, un yak et le wapiti fumer") == True
# Tout dans l'ordre
assert est_pangramme(ascii_lowercase) == True
# Tout à l'envers
assert est_pangramme(ascii_lowercase[::-1]) == True
# Tout en double
assert est_pangramme(ascii_lowercase * 2) == True
# Sans le "a"
assert est_pangramme(" ".join(ascii_lowercase[1:])) == False
# Sans le "b"
assert est_pangramme(" ".join("a" + ascii_lowercase[2:])) == False
# Sans le "z"
assert est_pangramme(" ".join(ascii_lowercase[:-1])) == False
### END HIDDEN TESTS
```

## Exercice 2 &ndash; Tri à bulles

#### Principe du tri à bulles

On parcourt le tableau de la fin vers le début, avec la variable d'indice `i`.  

Pour chaque parcours avec `i`, on parcourt le tableau avec un indice `j` en allant du début jusqu'à `i` (ou presque) ; si deux éléments consécutifs autour de l'indice `j` sont mal rangés, on les échange.

![illustration tri à bulles](tri_bulles.gif)

#### Travail à faire

Écrire une fonction telle que `tri_bulles(nombres)` opère un tri du tableau `nombres`.

- La fonction ne renvoie rien : inutile de placer return.
- Le tableau `nombres` est modifié par la fonction.


#### Exemples

```py
>>> nb_premiers = [2, 11, 3, 7, 5]
>>> tri_bulles(nb_premiers)
>>> nb_premiers
[2, 3, 5, 7, 11]
```

```py
>>> seul = [42]
>>> tri_bulles(seul)
>>> seul
[42]
```

```py
>>> vide = []
>>> tri_bulles(vide)
>>> vide
[]
```

<!-- grade_id: cell-2acb85ec4211863c, points: 0, grade: false, solution: true -->
```python
def tri_bulles(nombres):
    n = len(nombres)
    ### BEGIN SOLUTION
    for i in range(n - 1, -1, -1):
        for j in range(i):
            if nombres[j] > nombres[j + 1]:
                # nombres[j] et nombre[j + 1] à échanger
                nombres[j], nombres[j + 1] = nombres[j + 1], nombres[j]
    ### END SOLUTION
```

<!-- grade_id: cell-5b2b68f90926aeec, points: 4, grade: true, solution: false -->
```python
nb_premiers = [2, 11, 3, 7, 5]
nb_premiers_trie = [2, 3, 5, 7, 11]
tri_bulles(nb_premiers)
assert len(nb_premiers) == len(nb_premiers_trie), "Erreur, le tableau ne doit pas changer de taille"
for a, b, in zip(nb_premiers, nb_premiers_trie):
    assert a == b, "Erreur lors du tri de [2, 11, 3, 7, 5]"


seul = [42]
seul_trie = [42]
tri_bulles(seul)
assert len(seul) == 1, "Erreur, le tableau [42] ne doit pas changer de taille"
assert seul[0] == 42, "Erreur, un élément seul est déjà trié"

### BEGIN HIDDEN TESTS
from random import sample
for i in range(10):
    nombres = list(sample(range(10**9), 100+i))
    attendu = sorted(nombres)
    tri_bulles(nombres)
    assert len(nombres) == len(attendu), "Erreur, le tableau ne doit pas changer de taille"
    for a, b, in zip(nombres, attendu):
        assert a == b, "Erreur lors du tri"
### END HIDDEN TESTS
```

## Exercice 3 &ndash; Insertion dans une liste triée

#### Introduction

On souhaite écrire une fonction `insere` qui prend en argument un entier `a` et une liste `nombres` d'entiers triés par ordre croissant. Cette fonction insère la valeur `a` dans la liste de sorte que la liste reste triée.


```pycon
>>> exemple_1 = [1, 2, 4, 5]
>>> insere(3, exemple_1)
>>> exemple_1
[1, 2, 3, 4, 5]
```

```pycon
>>> exemple_2 = [1, 2, 7, 12, 14, 25]
>>> insere(7, exemple_2)
>>> exemple_2
[1, 2, 7, 7, 12, 14, 25]
```

```pycon
>>> exemple_3 = [2, 3, 4]
>>> insere(1, exemple_3)
>>> exemple_3
[1, 2, 3, 4]
```

#### Travail à faire

Une partie de la fonction `insere` est donnée ci-dessous :

```py
def insere(a, nombres):
    nombres.append(a)
    i = ...
    while (a < ...) and (i >= 0):
        nombres[i + 1] = ...
        nombres[i] = a
        i = ...
```

Écrire, **dans la cellule suivante**, le code complet de la fonction `insere`.

<!-- grade_id: cell-ef13caac38ecc02a, points: 0, grade: false, solution: true -->
```python
def insere(a, nombres):
    ### BEGIN SOLUTION
    nombres.append(a)
    i = len(nombres) - 2
    while (a < nombres[i]) and (i >= 0):
        nombres[i + 1] = nombres[i]
        nombres[i] = a
        i = i - 1
    ### END SOLUTION
```

<!-- grade_id: cell-8f641a9555973f64, points: 2, grade: true, solution: false -->
```python
# tests
exemple_1 = [1, 2, 4, 5]
insere(3, exemple_1)
assert exemple_1 == [1, 2, 3, 4, 5]

exemple_2 = [1, 2, 7, 12, 14, 25]
insere(7, exemple_2)
assert exemple_2 == [1, 2, 7, 7, 12, 14, 25]

exemple_3 = [2, 3, 4]
insere(1, exemple_3)
assert exemple_3 == [1, 2, 3, 4]

### BEGIN HIDDEN TESTS
nombres = []

insere(-101, nombres)
assert nombres == [-101], "Erreur insertion dans liste vide"

insere(11, nombres)
assert nombres == [-101, 11], "Erreur insertion singleton"

insere(0, nombres)
assert nombres == [-101, 0, 11], "Erreur insertion élément nul"

insere(-200, nombres)
assert nombres == [-200, -101, 0, 11], "Erreur insertion à gauche"

insere(200, nombres)
assert nombres == [-200, -101, 0, 11, 200], "Erreur insertion à droite"
### END HIDDEN TESTS
```