# Devoir maison 24 &ndash; $k$-somme maximale

## Somme maximale de k termes consécutifs

Écrire une fonction telle que `somme_maxi(valeurs, k)` renvoie la somme maximale de `k` termes consécutifs dans une liste de `valeurs`.

> On garantit que `valeurs` est de taille au moins égale à `k` et que `k > 0`.

#### Exemples

```pycon
>>> somme_maxi([0, 1, 2, 3, 2, 1, 0], 3)
7
>>> somme_maxi([0, 1, 2, 3, 2, 1, 0], 1)
3
```

#### Indice

On pourra commencer par faire le cumul des `k` premières valeurs pour initialiser une variable `maxi`.

On pourra ensuite faire une boucle qui ajoute la valeur suivante et retranche la première valeur.

## À vous de jouer !

<!-- grade_id: cell-55dadcb7ac03cdad, points: 0, grade: false, solution: true -->
```python
def somme_maxi(valeurs, k):
    k_somme = 0
    ### BEGIN SOLUTION
    for i in range(k):
        k_somme += valeurs[i]
    maxi = k_somme
    for i in range(k, len(valeurs)):
        k_somme += valeurs[i] - valeurs[i - k]
        if k_somme > maxi:
            maxi = k_somme
    return maxi
    ### END SOLUTION
```

<!-- grade_id: cell-ea7d74e8a851c880, points: 4, grade: true, solution: false -->
```python
assert somme_maxi([0, 1, 2, 3, 2, 1, 0], 3) == 7
assert somme_maxi([0, 1, 2, 3, 2, 1, 0], 1) == 3
### BEGIN HIDDEN TESTS
valeurs = [-42] * 1000
assert somme_maxi(valeurs, 100) == -42 * 100, "Erreur avec tableau constant"

assert somme_maxi(list(range(1000)), 100) == sum(range(900, 1000))
assert somme_maxi(list(range(-1000, 0)), 100) == -sum(range(101))
assert somme_maxi(list(range(1000, 0, -1)), 100) == sum(range(901, 1001))
assert somme_maxi(list(range(0, -1000, -1)), 100) == -sum(range(100))
### END HIDDEN TESTS
```